# The Game

*If you haven't watched the video about playing 20 questions against a computer using the Tree to help the computer ask you questions, this assignment probably won't make sense. Finish the lessons and come back.*

To make the 20 questions game work correctly, you will need to add the functionality of interacting with the provided btree class, followed by the logic of the game itself. This assignment is slightly different from earlier assignments as we are not dealing with input files, and now have user input and output. 

To make the notebook easier to understand, each portion is labeled as a new section. All of the code for interacting with the question/answer tree comes first, and each section has autograded tests. The code of these tests are generally hidden to keep from confusing you, as they interact with the tree in a different way than you should interact with it (this is the nature of testing code with data structures in Python). The comments in each cell tell you what the test is testing.

You may add cells to experiment with your functions wherever you see fit. But the requested functions must be filled in inside the provided cell.

For testing your final game, you have access to two saved game files: `/data/CS703/animalgame.txt` and `/data/CS703/plantgame.txt`. You can use these as input to your game, or you can create your own if you would like.

## The Overall process of the game

The very last code you'll write is the code to play the game, which will use the functions you write above it. You should read over that section first to make sure you understand what you are trying to do. The overall flow of the game is essentially the following steps:

1. Read in file containing pre-order traversal of game tree
2. Build BTree tree to represent game from this data
3. Play the game as many times as the user wants
  1. Tell them what game they are playing
  2. While you haven't reached an answer, keep asking questions and moving through the tree based on the user's yes/no answers
  3. Once you are at an answer, guess that answer by outputting it to the user
  4. Ask the use if you are correct
  5. If you are wrong, it's time to update the tree:
    1. Ask the user what the answer should have been
    2. Ask for a question that would have gotten you that answer from where you were in the tree, such that the correct answer was the "yes" answer and the answer you had guess was the "no" answer
    3. Update tree with this new question/answer
  6. Reset the game and ask the user if they want to play again

## What You'll Write

Due to the autograder tests the notebook looks long, but there's only a few cells where you are writing any code. The overall program is actually fairly short. You will be writing the following code in this notebook:

1. Setup function: This function takes the game in a pre-order list (NOTE: file reading happens elsewhere), and turns it into a BTree following the algorithm provided in the lessons.
2. Next_node function: This function returns the next node for consideration as the game is being played.
3. add_object: This function allows you to modify the tree that is being used as the game is played, when the computer guesses the wrong answer in the game (step 3.E above)
4. Code for reading in the file and storing it into the name variable and the pre-order list
5. Complete the code for playing the game. There are comments to guide you through the above steps, and some code is provided for you.

<!-- grade_id: cell-88a1ccc8febcbfc7, points: 0, grade: false, solution: true -->
```python
#import the ADTs you will need
from stack import Stack
from btree import BTree
```

## Loading 

Your first task is to create the function to load/start a game. You should create functions that act as defined by the doc strings provided. Note that it is up to you whether left or right is yes, but you should be consistent. You should use our discussion from class to help you write the setup function. If your algorithm is correct, you should need one line of code for every step in the algorithm; don't make it more complicated than it is.

<!-- grade_id: cell-719cc8c854e211e9, points: 0, grade: false, solution: true -->
```python
def setup(universe_list=None):
    '''Build a 20 questions game from a preorder traversal list, 
    using a stack from the provided stack class.
    Parameters are:
    * universe_list: a list of data in perorder order, to be loaded into tree.
    Returns:
    * if universe_list is not empty: the tree that contains all of the nodes in the correct 
    structure as defined by the universe_list parameter. 
    * otherwise: an empty tree with 1 node, no data.
    '''
    tree=BTree()
    if (universe_list is not None and len(universe_list)>0):
        my_stack=Stack()
        root_node = universe_list[0]
        tree.data = root_node
        my_stack.push(tree)
        for item in range(1, len(universe_list)):
            curr_node = my_stack.peek()
            new_node = BTree(universe_list[item])
            if curr_node.left == None:
                curr_node.left = new_node              
            else:
                if curr_node.right == None:
                    curr_node.right = new_node  
                    my_stack.pop()
            if '?' in universe_list[item]:
                my_stack.push(new_node)                

                
       
    return tree
```

```python
st=Stack()
st.push(3)

```

Output:
```
3
```

<!-- grade_id: cell-6adb65a93cadc713, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#You should be able to create a game with no data in it

tree=setup()
assert(tree.data is None)
assert(tree is tree)

```

```python
universe=["Does it produce something edible?",
          "Is it tropical?",
          "Is it spiky?",
          "Pineapple",
          "Banana Tree",
          "Apple Tree",
          "White Oak"]

test=setup(universe)
test.print_tree()
```

Output:
```
-> White Oak
-> Does it produce something edible?
        -> Apple Tree
    -> Is it tropical?
            -> Banana Tree
        -> Is it spiky?
            -> Pineapple
```

<!-- grade_id: cell-48e5b8158450277a, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Creating a Game with an empty universe should work, but provide an empty tree.

empty=[]

test_tree=setup(empty)
assert(test_tree.data is None)
assert(test_tree.left is None and test_tree.right is None)
```

<!-- grade_id: cell-6455ee7a464bc7e1, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test that we can create a Game from a universe

universe=["Does it produce something edible?",
          "Is it tropical?",
          "Is it spiky?",
          "Pineapple",
          "Banana Tree",
          "Apple Tree",
          "White Oak"]

test=setup(universe)
test.print_tree()
assert(test.data == "Does it produce something edible?"), "root node isn't correct"
```

Output:
```
-> White Oak
-> Does it produce something edible?
        -> Apple Tree
    -> Is it tropical?
            -> Banana Tree
        -> Is it spiky?
            -> Pineapple
```

<!-- grade_id: cell-81ebe4df22368b0a, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test that when we create from an existing list whether the children are populated

universe=["Does it produce something edible?",
          "Is it tropical?",
          "Is it spiky?",
          "Pineapple",
          "Banana Tree",
          "Apple Tree",
          "White Oak"]

test=setup(universe)

c=test.left
assert(c.data == "White Oak" or c.data=="Is it tropical?"), "Left child is not added correctly"
```

<!-- grade_id: cell-be613a63330d43c6, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test that when we create from an existing list whether the children are BOTH populated with the correct options

universe=["Does it produce something edible?",
          "Is it tropical?",
          "Is it spiky?",
          "Pineapple",
          "Banana Tree",
          "Apple Tree",
          "White Oak"]

test=setup(universe)

root=test
c=test.left
prev_answer=(c.data == "White Oak")

c=root
c=c.right
assert(((not prev_answer) and (c.data == "White Oak")) or 
       (prev_answer and (c.data == "Is it tropical?"))), "Both children are not correct"
```

<!-- grade_id: cell-aa51b3be8ddd101b, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test that when we create from an existing list whether an arbitrary element can be reached correctly.

universe=["Does it produce something edible?",
          "Is it tropical?",
          "Is it spiky?",
          "Pineapple",
          "Banana Tree",
          "Apple Tree",
          "White Oak"]

test=setup(universe)

# Note that this code is for testing purposes and is not an example of code you should be writing
# to solve any of the problems you need to solve in this assignment
root=test
curr=test.left
left_is_no=(curr.data == "White Oak")

def test_guess(left_is_no,guess,tree):
    if((left_is_no and guess=="Y") or ((not left_is_no) and guess=="N")):
        return tree.right
    else:
        return tree.left
       
test=root
test=test_guess(left_is_no, "Y", test)
test=test_guess(left_is_no, "Y", test)
test=test_guess(left_is_no, "N", test)
assert(test.data == "Banana Tree"), "Could not navigate to Banana tree correctly"

del test_guess

```

## Moving to Next Node

The game is going to guess what the user is thinking of, using its tree of data. We need a function that will return to us the next step in our guessing game based on the information the user has given us. For instance, if we are currently looking at a question and the user has told us the answer is "no", we should return whatever the next step should be in the "no" situation.

<!-- grade_id: cell-a292e650f0399ffe, points: 0, grade: false, solution: true -->
```python
def next_node(curr_node,yes_or_no):
    '''Updates curr_node to be the next node to look at, depending on if they said yes or no.
        yes_or_no holds their guess of y/n
        curr_node tells us which node of the tree is the question we are currently guessing on.
        This function should return the next node in the game.'''
    if yes_or_no == 'Y' or yes_or_no =='yes' or yes_or_no == 'y':
        curr_node = curr_node.left
        return curr_node
    elif yes_or_no == 'N' or yes_or_no =='no' or yes_or_no == 'n':
        curr_node = curr_node.right
        return curr_node
```

```python
# test=setup()
# test.data="Question?"
# test.left=BTree("yes")
# test.right=BTree("no")
# curr=test
# curr=next_node(curr,"Y")
# print(curr.data)


test=setup()
test.data="Question?"
test.left=BTree("yes")
test.right=BTree("no")
curr = test
curr = next_node(curr,"N")
print(curr.data)
```

Output:
```
no
```

<!-- grade_id: cell-88951eae238d4699, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test that guess does not adjust the root of the tree, but does move the current node either left of right.
# Code for testing should not be used as an example of how to write your cod

test=setup()
test.data="Question?"
test.left=BTree("yes")
test.right=BTree("no")
curr=test
curr=next_node(curr,"Y")
print(curr.data)
assert(test.data == "Question?")
assert(curr.data == "yes" or curr.data == "no")
```

Output:
```
yes
```

<!-- grade_id: cell-e3b5ca713658d9eb, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
# Test that guessing Y moves you in a different direction from guessing N
# Code for testing should not be used as an example of how to write your code

test=setup()
test.data="Question?"
test.left=BTree("yes")
test.right=BTree("no")
curr=next_node(test,"Y")
prev_answer=(curr.data == "yes")

test=setup()
test.data="Question?"
test.left=BTree("yes")
test.right=BTree("no")
curr=next_node(test,"N")

assert((prev_answer and (curr.data == "no")) or ((not prev_answer) and (curr.data == "yes"))), "yes and no are not in different directions"
```

## Modifying the Tree 

In addition to playing the guessing game, we can also increase our knowledge. This function allows us to add onto our tree at our current location. We would only use this function while playing the game.

Recall that we need to make sure it is clear when a node is a question node. We need to add our new question node to the tree at the current location, then attach the new yes item, and either a new no item or the old no item we are extending on. Drawing things out on paper can help.

<!-- grade_id: cell-b1e87ae3587f715b, points: 0, grade: false, solution: true -->
```python
# this function lets us add on to our tree

def add_object(curr_node,question,yes_thing,no_thing=None):
    '''Add a new thing to our Universe. If no_thing is None, assume that the current node was an incorrect
    guess and is thus the "No" answer to the question.
    Parameters are:
    * curr_node: the node of the tree where the question is to be added.
    * question: The question that is to be put at this location.
    * yes_thing: The item in the yes child
    * no_thing: The item in the no child; if None, then the curr_node's data is the question'''

    saved = curr_node.data
    if '?' in question:
        curr_node.data = question
        if no_thing:
            curr_node.left = BTree(yes_thing)
            curr_node.right = BTree(no_thing)
        else:
            curr_node.right = BTree(yes_thing)
            curr_node.left = BTree(saved)

```

<!-- grade_id: cell-f89d1d4d96ffa170, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test whether add_object works on an empty tree

test=setup()
add_object(test,"Question?", "yes", "no")
test.print_tree()
assert(test.data == "Question?"), "question was added as the root"
assert(test.left.data == "yes" or test.right.data == "yes"), "yes was added as a child"
assert(test.left.data == "no" or test.right.data == "no"), "no was added as a child"
```

Output:
```
-> no
-> Question?
    -> yes
```

<!-- grade_id: cell-28d0ba95528565f7, points: 0, grade: true, solution: false -->
```python
# AUTOGRADING TEST
#Test that add_opbject works as intended with a single parameter

test=setup()
test.data="Question?"
test.left=BTree("yes")
test.right=BTree("no")
curr=test.left
add_object(curr,"New Question?", "new yes")
test.print_tree()
assert(test.left.data == "New Question?"), "Did not correctly add new question"
assert((test.left.left.data == "yes" and test.left.right.data == "new yes") or \
         (test.left.right.data == "yes" and test.left.left.data == "new yes")), "did not correct add a new question node with new option"  
```

Output:
```
-> no
-> Question?
        -> new yes
    -> New Question?
        -> yes
```

The cell below is for you to experiment with your Game functions, and will not be graded.

```python

```

# Time to Play the Game
## Game Step 1: Load the game file data into a list
The next cell sets up the game by requesting a game file and loading it into the initial_universe list. Note that the first element in the file is actually the name of the game, not part of the game, so should be saved into the universe_name variable instead of the list.

```python
# Initial setup of the game

print("Welcome to the Jumbo Toys 20 Questions game.", 
      "This game will load a universe of objects and",
      "then ask you yes or no questions to figure out",
      "what you are thinking about!\n\n")

universe_file=input("What universe would you like to load? ")
```

Output:
```
Welcome to the Jumbo Toys 20 Questions game. This game will load a universe of objects and then ask you yes or no questions to figure out what you are thinking about!


What universe would you like to load? /data/CS703/animalgame.txt
```

<!-- grade_id: cell-a2610a58333c8558, points: 0, grade: false, solution: true -->
```python
# Open the file and read it into the list with basic error handling try/except

universe_name = ""
initial_universe = []

with open('/data/CS703/animalgame.txt', "r") as animalGame:
    for line in animalGame:
        line = line.strip('\n')     
        line.split(',')   
        initial_universe.append(line)
    if '?' not in initial_universe[0]:
        initial_universe.remove(initial_universe[0])
```

```python
# tests to check file was read correctly
print(universe_name)
print(initial_universe)#debugging for grading
```

Output:
```
['Is it a mammal?', 'Does it eat meat?', 'Does it have stripes?', 'Tiger', 'Does it have a domesticated cousin?', 'Wolf', 'Bear', 'Does it have tusks?', 'Elephant', 'Does it have a long neck?', 'Giraffe', 'Does it live in the water?', 'Whale', 'Rabbit', 'Does it have legs?', 'Does it have teeth?', 'Crocodile', 'Frog', 'Does it live in the water?', 'Fish', 'Snake']
```

```python
print(len(initial_universe))
```

Output:
```
21
```

## Game Step 2: The actual game

Now we finally have everything set up, and we can play the game!

Pay attention to the comments that tell you want you need to add to the loop. You have 4 steps to add in:

1. Tell the user the name of the universe
2. Play the game
3. Add in new information if the game guessed the wrong item
4. reset the game once it's over so that it can be played again.

### User Interface 

I have provided a function that will help you take in yes or no input from the user in a safe way. 

```python
def get_y_or_n(inputstring):
    '''Returns y or n (lowercase) according to user input. Keeps prompting if y or n is not given.
    Input string is the intitial prompt for the user.'''
    print("",end="")
    return_string=input(inputstring)[0:1].lower()
    while(return_string != "y" and return_string != "n"):
        return_string=input("Please answer Y or N: ")[0:1].lower()
        
    return return_string
```

<!-- grade_id: cell-f37858e27905b6af, points: 0, grade: false, solution: true -->
```python
#Initialize the game by calling setup
my_game=setup(initial_universe)
current = my_game # keep track of where we are in the game tree

# Keep playing the game as long as the user is still interested in playing
playagain="y"
while(playagain=="y"):
    # Tell user what universe they are guessing from
    print("Think of an animal and I will try to guess it!")
 
    # Actually play the game
    # Until you've hit an answer, keep asking the user questions and moving 
    # through the tree based on their answers.
    while '?' in current.data:
        print(current.data)
        YN = get_y_or_n("Y/N: ")
        current = next_node(current, YN) 
       
    # Once you've hit an answer, output what you think it is.
    if '?' not in current.data:
        print("I think I got it!", current.data,"?")
    
    userin = get_y_or_n("Am I right? ")
    if(userin == "n"):
#         #Oops, the game guessed wrong. Add a new item at this spot in the tree, getting info from user!
        new_item = input("Oh no! What is it?").strip()
        question = input('Please enter in a question where '+ current.data+ ' would be No and '+new_item+ ' would be yes: ')
        add_object(current, question, new_item, current.data)
    else:
        print("I guessed it!")
        
    #Get ready to play again by resetting the game
    current = my_game
    
    playagain = get_y_or_n("\nPlay again? ")
```

Output:
```
Think of an animal and I will try to guess it!
Is it a mammal?
Y/N: y
Does it eat meat?
Y/N: y
Does it have stripes?
Y/N: y
I think I got it! Tiger ?
Am I right? n
Oh no! What is it?taz
Please enter in a question where Tiger would be No and taz would be yes: does it live in australia?

Play again? y
Think of an animal and I will try to guess it!
Is it a mammal?
Y/N: y
Does it eat meat?
Y/N: y
Does it have stripes?
Y/N: y
does it live in australia?
Y/N: y
I think I got it! taz ?
Am I right? y
I guessed it!
```

```python
my_game.print_tree()
```

# Grading

The cell below is to hold your grade on the assignment. The number has no meaning, but is used to represent the letter:

* 3: E
* 2: M
* 1: R
* 0: N

<!-- grade_id: cell-244439e81064d90f, points: 3, grade: true, solution: true -->
```python

```

```python

```

```python

```

```python

```